---
title         : Data wrangling
description   : Introduces split-apply-combine
free_preview  : true

--- type:NormalExercise lang:r xp:100 skills:1 key:1b4b89b117
## Casting data
In addition to making plots, R has tools to help you prepare tabular views of your data. This task can
be viewed as a two-part process: first, selecting (aka subsetting) the data points that you want to present,
and then arranging that data into the rows and columns that you need.
*** =instructions



*** =hint

no hints yet

*** =pre_exercise_code
```{r}


```

*** =sample_code
```{r}




```

*** =solution
```{r}



```

*** =sct
```{r}
test_error()
```